```python
# Initialize Otter
import otter
grader = otter.Notebook("worksheet3.ipynb")
```

# Worksheet 3 - Advanced Wrangling With Pandas

These exercises complement DSCI 511 lecture 3.

## Exercises

```python
import pandas as pd
import numpy as np
import datetime as dt
```

### Q1

rubric={autograde:1}

Given the following variables:
```
thing = "light"
speed = 299792458 
```
Use f-strings to print:
```
The speed of light is 299792458 m/s.
```

```python
thing = "light"
speed = 299792458  # m/s
q5_answer = ...
print(q5_answer)
```

```python
grader.check("q1")
```

### Q2

rubric={autograde:1}

In this set of practice exercises we'll be looking at a cool dataset of real passwords (made available from actual data breaches) sourced and compiled from [Information is Beautiful](https://informationisbeautiful.net/visualizations/top-500-passwords-visualized/?utm_content=buffer994fa&utm_medium=social&utm_source=twitter.com&utm_campaign=buffer) and contributed to [R's Tidy Tuesday project](https://github.com/rfordatascience/tidytuesday). These passwords are common ("bad") passwords that you should avoid using!

The dataset has the following columns:

|variable          |class     |description |
|:-----------------|:---------|:-----------|
|rank              |int    | popularity in their database of released passwords |
|password          |str | Actual text of the password |
|category          |str | What category does the password fall in to?|
|value             |float    | Time to crack by online guessing |
|time_unit         |str | Time unit to match with value |
|offline_crack_sec |float    | Time to crack offline in seconds |
|rank_alt          |int    | Rank 2 |
|strength          |int    | Strength = quality of password where 10 is highest, 1 is lowest, please note that these are relative to these generally bad passwords |
|font_size         |int    | Used to create the graphic for KIB |


In these exercises, we're only interested in the `password`, `value` and `time_unit` columns so import only these columns as a dataframe named `df` from this url: <https://raw.githubusercontent.com/rfordatascience/tidytuesday/master/data/2020/2020-01-14/passwords.csv>

Note: There are several 'NA' rows at the end. Pass the optional parameters `skipfooter=7` and `engine=python` in your call to `pd.read_csv` in order to skip these.

```python
...
```

```python
grader.check("q2")
```

### Q 3.

rubric = {autograde:1}

How many password contain the sequence `123`? Store your output as a single value in `q3` object

```python
q3 = ...
```

```python
grader.check("q3")
```

### Q 4. 

rubric = {autograde:1}


The `value` column contains times in many different units. The code in the cell below will convert them all to hours so that they can be easily compared. Run the cell below before proceeding.

```python
# RUN THIS CELL

units = {
    "seconds": 1 / 3600,
    "minutes": 1 / 60,
    "days": 24,
    "weeks": 168,
    "months": 720, # assuming there are 30days in each month
    "years": 8760,
}

for key, val in units.items():
    df.loc[df['time_unit'] == key, 'value'] *= val 

df['time_unit'] = 'hours'
df.head()
```

#### Question:

What is the average time in hours needed to crack these passwords that begin with `123`? 

Store your output as a single value in `avg_crack_123`

```python
avg_crack_123 = ...
```

```python
grader.check("q4")
```

### Q5.

Create a `DatetimeIndex` called `date_range` consisting of dates starting 2010-01-01 00:00:00 until today, with daily frequency.

Hint: To get today's datetime, you can use `pd.Timestamp.now()`

```python
date_range = None
...
```

```python
grader.check("q5")
```

### Q 6

The code cell below creates a new data frame called `df2` using the `date_range` you defined in the previous question. Run this cell.

```python
# create a DataFrame with the datetime index
df2 = pd.DataFrame(index=date_range)

df2
```

For this question, create the following 4 additional columns for the dataframe `df2`:
- year
- month_name
- day
- weekday_name.

Fill them in with the appropriate values based on the index.

```python
...
```

```python
grader.check("q6")
```

Tests for q1 (points: 1)
```
>>> assert q5_answer == 'The speed of light is 299792458 m/s.'
```


Tests for q2 (points: 1)
```
>>> assert type(df) == pd.core.frame.DataFrame
>>> assert 'password' and 'value' and ('time_unit' in df.columns)
```


Tests for q3 (points: 1)
```
>>> assert q3 == 12
```


Tests for q4 (points: 1)
```
>>> assert np.isclose(avg_crack_123, 2052.877493287037, atol=0.01)
```


Tests for q5 (points: 1)
```
>>> assert type(date_range) == pd.core.indexes.datetimes.DatetimeIndex
>>> assert date_range.freq == 'D'
>>> assert date_range[0] == dt.datetime(2010, 1, 1, 0, 0)
```


Tests for q6 (points: 1)
```
>>> assert 'year' and 'month_name' and 'day' and ('weekday_name' in df2.columns)
>>> assert df2['year'].min() == 2010 and df2['year'].max() == 2024
>>> assert df2['day'].min() == 1 and df2['day'].max() == 31
>>> assert df2['month_name'].unique().tolist() == ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December']
>>> assert df2['weekday_name'].unique().tolist() == ['Friday', 'Saturday', 'Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday']
```
