```python
# Initialize Otter
import otter
grader = otter.Notebook("worksheet8.ipynb")
```

# Worksheet 8 - Classes and Style

These exercises complement DSCI 511 lecture 8.

## Exercises

```python
import pandas as pd
import numpy as np
import math
```

### Q1: Creating a Class

rubric={autograde:1}

Create a class called `Circle`. It should have the following characteristics:
1. It should be initiated with the argument `radius` and store this as an instance attribute.
2. Have a method `area()` which calculates the area of the circle. $A = \pi * r^2$
3. Have a method `circumference()` which calculates the circumference of the circle. $C = 2*\pi*r$
4. Have the method `__str__()` which is a magic method that controls what is output to the screen when you `print()` an instance of your class (learn more [here](https://realpython.com/lessons/how-and-when-use-__str__/)). This method should return the string `"A circle with radius r"` where `r` is the radius value. For example:
```python
c1 = Circle(1)
print(c1) # EXPECTED OUTPUT: A circle with radius 1
```

```python
...
```

```python
grader.check("q1")
```

### Q2: Class Inheritance

rubric={autograde:1}

**Instructions**

1. Define a class called `Animal` with an `__init__()` method that takes a `name` parameter and sets it as an instance attribute.
2. Define a class called `Dog` that inherits from the `Animal` class.
3. In the `Dog` class, define an `__init__()` method that takes a `name` and `breed` parameter, and calls the `__init__()` method of the `Animal` class using `super()`.
4. In the `Dog` class, define a `bark()` method that returns the string "Woof!" (Don't print "Woof!")

```python
...
```

```python
grader.check("q2")
```

### Q 3.

rubric = {autograde:1}

This exercise is built upon your answer in Q3.

- Add a `count` **class variable** to the `Dog` class. The `__init__()` method increments the `count` variable each time a new instance of the class is created.

- Add a `get_dog_count()` **class method** to the `Dog` class that returns the number of `Dog` instances that have been created.

- Add a `reset_dog_count()` **class method** to the `Dog` class that sets the count to zero

```python
...
```

```python
# Some tests to check your code
Dog.reset_dog_count()
print(f"The number of dogs after resetting is {Dog.get_dog_count()}")
      
dog1 = Dog("Fido", "Golden Retriever")
print(f"After create a dog1 instance, the number of dogs is {Dog.get_dog_count()}")

dog2 = Dog("Hunky", "German Shepherd")
print(f"After create a dog2 instance, the number of dogs is {Dog.get_dog_count()}")

print(f"Get the count of dogs: There are {Dog.get_dog_count()} dogs") # Should return 2
```

```python
grader.check("q3")
```

### Q 4. 

rubric = {manual:0}

Fix the following code to make sure it aligns with PEP8 style guide.

```python
def q1_incorrect():
    CarModel='Toyota'
    Make="RAV4"
    return CarModel + Make
```

<!-- BEGIN QUESTION -->



```python
def q1():
    ...
```

<!-- END QUESTION -->

### Q5

rubric = {manual:0}

Fix the following code to make sure it aligns with PEP8 style guide

def q2_incorrect():
    gpa = (course1 + course2 + course3 + course4 + course5 + course6 + course7 + course8 + course9 + course10) / 10

<!-- BEGIN QUESTION -->



```python
def q2():
    ...
```

<!-- END QUESTION -->

### Q 6

rubric = {manual:0}

Fix the following code to make sure it aligns with PEP8 style guide.

class car_model:
    def __init__(self, car_model, make):
        self.car_model = car_model
        self.make = make   
    def __str__(self):
        return f"{self.car_model} {self.make}"
class car_sale(car_model):
    def __init__(self, car_model, make, price):
        super().__init__(car_model, make)
        self.price = price   
    def sell(self):
        return f"{self.car_model} {self.make} sold for {self.price}"

<!-- BEGIN QUESTION -->



```python
...
```

<!-- END QUESTION -->

Congratulations! You are done the **LAST** worksheet!!! Pat yourself on the back, and submit your worksheet to Gradescope!

![](https://i.giphy.com/media/v1.Y2lkPTc5MGI3NjExbTljbW5uOTd6Y3p0cTl0Y2ZqbnI0aTM5MGNibXUzOTl0cTRobGhxciZlcD12MV9pbnRlcm5hbF9naWZfYnlfaWQmY3Q9Zw/gLWTWCg96RA0btDUm0/giphy.gif)

>### Q4 Solutions
>
>```
>def q1():
>    car_model = "Toyota"
>    make = "RAV4"
>    return car_model + make
>```
>
>
>### Q5 Solutions
>
>```
>def q2():
>    # BEGIN SOLUTION
>    gpa = (
>        course1
>        + course2
>        + course3
>        + course4
>        + course5
>        + course6 
>        + course7
>        + course8
>        + course9
>        + course10
>    ) / 10
>```
>
>
>### Q6 Solutions
>
>```
>class CarModel:
>    def __init__(self, car_model, make):
>        self.car_model = car_model
>        self.make = make
>
>    def __str__(self):
>        return f"{self.car_model} {self.make}"
>
>
>class CarSale(CarModel):
>    def __init__(self, car_model, make, price):
>        super().__init__(car_model, make)
>        self.price = price
>
>    def sell(self):
>        return f"{self.car_model} {self.make} sold for {self.price}"
>```

Tests for q1 (points: 1)
```
>>> assert hasattr(Circle(3), 'radius'), 'The Circle class needs a radius attribute.'
>>> assert Circle(3).radius == 3, 'Radius is not being set properly.'
>>> assert math.isclose(Circle(3).area(), 28.3, abs_tol=0.1), 'Area method does not calculate properly.'
>>> assert math.isclose(Circle(3).circumference(), 18.8, abs_tol=0.1), 'Circumference method does not calculate properly.'
>>> assert 'A circle with radius 3' in Circle(3).__str__(), 'The __str__ method does not return the correct string.'
```


Tests for q2 (points: 1)
```
>>> assert Animal('Miso').name == 'Miso', 'Animal name is not set properly'
>>> assert hasattr(Animal('Miso'), 'name'), 'Class Animal does not have attribute name'
>>> assert issubclass(Dog, Animal), 'Dog is not a subclass of Animal'
>>> dog1 = Dog('Chunky', 'German Shepherd')
>>> assert dog1.name == 'Chunky', 'Dog name is not set properly'
>>> assert dog1.breed == 'German Shepherd', 'Dog breed is not set properly'
>>> assert 'Woof!' in dog1.bark(), 'Dog bark method is not set properly'
```


Tests for q3 (points: 1)
```
>>> Dog.reset_dog_count()
>>> assert Dog.count == 0
>>> dog1 = Dog('Fido', 'Golden Retriever')
>>> assert Dog.count == 1
>>> dog2 = Dog('Buddy', 'Labrador Retriever')
>>> assert Dog.count == 2
>>> assert Dog.get_dog_count() == 2
```
