```python
# Initialize Otter
import otter
grader = otter.Notebook("worksheet4.ipynb")
```

## Worksheet 4 - Grouping

Run the cell below to import the necessary packages for this worksheet.

```python
import pandas as pd
import numpy as np
```

### Q1
rubric={autograde:1}

We return to the Narrabeen beach survey dataset that we encountered in Lab 1. Read in the data from `data/beach_data.csv` and save it to a data frame named `beach_df`, making sure to set the dates as index in a pandas datetime format.

```python
beach_df = ...
```

```python
grader.check("q1")
```

## Q2

rubric={autograde:1}

The data is quite irregularly spaced, with a frequency ranging between a few weeks to a few months. Your first tasks are to

* Resample `beach_df` from the previous step to monthly intervals with the mean as the aggregation function.

* After resampling, subtract the mean value of each resulting "location" column from the same column. This will help to see if a location on the beach is narrower (negative numbers) or wider (positive numbers) at a certain time compared to the average.

* Finally, reshape the data set so it has two columnns `location` and `width`. The `location` column should indicate which of the five locations each measurement was taken at.

The final data frame should look something like the following:


|   &nbsp;   |    location   |   width   |
| ---------  | ------        | ------    | 
|**datetime**|     &nbsp;    |  &nbsp;   |                   
|1980-01-01  |   location_1  |   4.796227|
|1980-02-01  |  location_1   |  4.296227 |
|1980-03-01  | location_1    | 0.596227  |
|1980-04-01  |location_1     |5.296227   |
|1980-05-01  |   location_1  | -13.903773|
|...         |      ...      |    ...    |
|2019-07-01  | location_5    |-8.395284  |
|2019-08-01  |location_5     |-7.578618  |
|2019-09-01  |   location_5  |-15.778618 |
|2019-10-01  |  location_5   |-16.478618 |
|2019-11-01  | location_5    |-16.895284 |




```python
...
```

```python
beach_df
```

```python
grader.check("q2")
```

## Q3

rubric={autograde: 1}

* Extract the month from the `DatetimeIndex`, and use it to create a new `month` column in the original `beach_df` dataframe.
* Group the data by month and aggregate it based on the mean. Use this to find the three months where the beach was narrowest on average. Store your answer in a dataframe called `beach_top_3` with 3 rows, 'month' as index and 1 column for 'width'.

```python
...
```

```python
beach_top_3
```

```python
grader.check("q3")
```

## Q4
rubric={autograde:1}

Perform a double `groupby()` to determine the combination of **month** *and* **location** for which the beach is the widest. This time, aggregate the data based on the median value per group.

Store your output in a dataframe called `beach_df_widest` with a single row. It should have the `month` and `location` as the index, and 1 column called `width`.

```python
beach_df_widest = ...
```

```python
grader.check("q4")
```

## Q5
rubric={autograde: 1}

Run the code cell below to create two dataframes, `dates` and `rooms`.

```python
# RUN THIS CELL

dates = pd.DataFrame(
                    { 'name': ['Kate', 'Kaiyun', 'Prajeet', 'Tiffany', 'Mohit', 'Eric'],
                      'day': ['Monday', 'Tuesday', 'Tuesday', 'Wednesday', 'Thursday', 'Wednesday'],
                      'time': ['5pm', '4:30pm', '1pm', '1pm', '4pm', '4pm']
                    }
                    )

rooms = pd.DataFrame(
                    {
                      'day': ['Wednesday', 'Wednesday', 'Thursday', 'Tuesday', 'Tuesday', 'Monday'],
                      'time': ['4pm', '1pm', '4pm', '4:30pm', '1pm', '5pm'],
                      'room': ['MCML 160', 'ESB 3174', 'ICCS X153', 'ICCS X153', 'ESB 1046', 'ICCS X153']
                    }
                    )
```

The data frame `dates` contains the dates and times for some instructor and TA office hours for DSCI 511 this term. In `rooms` you will find room booking information for these dates and times. Your task is to:

* Combine the two data frames meaningfully to make a 'time table' of office hours. You will have to decide whether `concat()` or `merge()` is more suitable for this.
* Save this time table in a data frame titled `oh_info` with index `name`
* Make sure `oh_info` has the three columns `day`, `time` and `room`.

```python
oh_info = None
...
```

```python
oh_info
```

```python
grader.check("q5")
```

Tests for q1 (points: 1)
```
>>> assert hasattr(beach_df.index, 'month'), 'Did you remember to parse the dates?'
>>> assert np.isclose(beach_df.loc['1993-01-11', 'location_3'], 76.0, atol=0.1)
```


Tests for q2 (points: 1)
```
>>> assert beach_df.shape == (2395, 2), 'Did you melt the data frame?'
>>> assert beach_df.sort_index().index[105] - beach_df.sort_index().index[100] == pd.Timedelta('30 days'), 'Did you resample to a monthly frequency?'
```


Tests for q3 (points: 1)
```
>>> assert beach_top_3.shape == (3, 1)
>>> assert 'month' in beach_top_3.index.names, "Error: DataFrame does not have index 'month'"
>>> assert 'width' in beach_top_3.columns, "Error: DataFrame does not have a column 'width'"
>>> assert np.isclose(beach_top_3.query('month == 7')['width'].values[0], -1.678812, atol=0.1)
>>> assert np.isclose(beach_top_3.query('month == 8')['width'].values[0], -1.766926, atol=0.1)
>>> assert np.isclose(beach_top_3.query('month == 9')['width'].values[0], -1.935821, atol=0.1)
```


Tests for q4 (points: 1)
```
>>> assert len(beach_df_widest) == 1
>>> assert 'month' in beach_df_widest.index.names
>>> assert 'location' in beach_df_widest.index.names
>>> assert 'width' in beach_df_widest.columns
>>> assert beach_df_widest.index[0] == (5, 'location_5')
>>> assert np.isclose(beach_df_widest['width'].values[0], 10.471382, atol=0.1)
```


Tests for q5 (points: 1)
```
>>> assert oh_info.shape == (6, 3)
>>> assert oh_info.loc['Prajeet', 'room'] == 'ESB 1046'
>>> assert oh_info.loc['Mohit', 'day'] == 'Thursday'
```
