```python
# Initialize Otter
import otter
grader = otter.Notebook("worksheet7.ipynb")
```

# Worksheet 7 - Functions & Testing

These exercises complement [DSCI 511 lecture 7](../lecture7-functions-testing.ipynb).

```python
import pytest
import math
import unittest
import io
from contextlib import redirect_stdout
```

## 1. Functions

### 1.1 Website domain

Create a function `website()` that grabs the website domain from a url string. For example, if your function is passed `"www.google.com"`, it should return `"google"`.

Hint: You can use the `split()` built-in function to parse a url

```python
def website(url):
    ...
```

```python
# TEST CASE
website('www.apple.com') # EXPECT OUTPUT: 'apple'
```

```python
grader.check("q1.1")
```

### 1.2 Divisible?

Create a function `divisible(a, b)` that accepts two integers (`a` and `b`) and returns `True` if `a` is divisble by `b` without a remainder. For example, `divisible(10, 3)` should return `False`, while `divisible(6, 3)` should return `True`.

```python
def divisible(a, b):
    ...

print(divisible(10, 3)) # output should be FALSE
print(divisible(6, 3)) # output should be TRUE
```

```python
grader.check("q1.2")
```

### 1.3 Generators

This question is a little harder. Create a generator function called `listgen(n)` that yields numbers from 0 to n, in batches of lists of maximum 10 numbers at a time. For example, your function should behave as follows:

```python
for batch in listgen(25):
    print(batch)
    
[0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
[10, 11, 12, 13, 14, 15, 16, 17, 18, 19]
[20, 21, 22, 23, 24]
```

```python
def listgen(n):
    ...

for batch in listgen(25):
    print(batch)
```

```python
grader.check("q1.3")
```

### 1.4 *args and **kwargs

Create a function `lucky_sum()` that takes all the integers a user enters and returns their sum. *However*, if one of the values is 13 then it does not count towards the sum, nor do any values to its right.

For example, your function should behave as follows:

```python
lucky_sum(1, 2, 3, 4)
10

lucky_sum(1, 13, 3, 4)
1

lucky_sum(13)
0
```

*This example is inspired by the related [codingbat challenge](https://codingbat.com/prob/p130788).*

```python
def lucky_sum(*args):
    ...
```

```python
grader.check("q1.4")
```

### 2.1 Unit tests

rubric = {autograde: 1}

The function `area()` accepts the argument `radius` and calculates the area of a circle. In the function `unit_test_area()` below, write three tests using `assert` statements for the following conditions:

1. Check if the function `area()` returns the correct data type
2. Check if the function `area()` returns the correct calculation up to 2 decimals with a input value of your choice
3. Check if the function `area()` returns the correct calculation up to 2 decimals with a different input value

(hint: `math.isclose(..., abs_tol=0.01)`)

```python
def area(radius):
    """Calculate the area of a circle based on the given radius."""
    return math.pi * radius ** 2
    
def unit_test_area():
    ...

```

```python
grader.check("q2.1")
```

### 2.2 EAFP unit test

rubric = {autograde: 1}

In the spirit of the EAFP (easier to ask for forgiveness than permission) philosophy. Modify the code of the function `area()` and add a `try`/`except` statement to catch the `TypeError` if users pass a non-numeric input to `area()`. Print out a custom message "radius should be a number but you entered a X" where X is the data type of the input object. For example:

```python
area('10')  # EXPECTED OUTPUT: radius should be a number but you entered a <class 'str'>
```

Note: You should print out an error message instead of raising an error

```python
def area(radius):
    ...

```

```python
grader.check("q2.2")
```

### 2.3 LBYL unit test

rubric = {autograde: 1}

In the spirit of the LBYL (look before you leap) philosophy. Modify the code of the function `area()` and add a conditional `if`/`else` statement to make sure that a user has passed a number (`int` or `float`) to the `area()` function. If they pass something else, raise a `TypeError` with a message of your choice. For example:

```python
area('10') # EXPECTED OUTPUT: TypeError: Your input is not a number!
```

Note: You should raise an error instead of printing out an error message.

**Hint:** Use the built-in Python function `isinstance(obj, (class1, class2, ...)` to inspect input object types.

```python
def area(radius):
    ...
```

```python
grader.check("q2.3")
```

Tests for q1.1 (points: 1)
```
>>> assert website('www.bbc.com') == 'bbc'
>>> assert website('www.amazon.ca') == 'amazon'
>>> assert website('www.yahoo.co.jp') == 'yahoo'
```


Tests for q1.2 (points: 1)
```
>>> assert divisible(9, 3)
>>> assert not divisible(16, 3)
```


Tests for q1.3 (points: 1)
```
>>> gen = listgen(25)
>>> assert next(gen) == [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
>>> assert next(gen) == [10, 11, 12, 13, 14, 15, 16, 17, 18, 19]
>>> assert next(gen) == [20, 21, 22, 23, 24]
```


Tests for q1.4 (points: 1)
```
>>> assert lucky_sum(1, 2, 3) == 6
>>> assert lucky_sum(1, 2, 13) == 3
>>> assert lucky_sum(1, 2, 13, 4, 5) == 3
>>> assert lucky_sum(13, 2, 3, 4) == 0
```


Tests for q2.1 (points: 1)
```
>>> def area(radius):
...     return radius ** 5 / 2
>>> 
>>> def test_value_error():
...     with pytest.raises(Exception):
...         unit_test_area()
>>> test_value_error()
```
```
>>> def area(radius):
...     return 'string'
>>> 
>>> def test_value_error():
...     with pytest.raises(Exception):
...         unit_test_area()
>>> test_value_error()
```


Tests for q2.2 (points: 1)
```
>>> f = io.StringIO()
>>> with redirect_stdout(f):
...     area([10])
>>> result = f.getvalue()
>>> assert "radius should be a number but you entered a <class 'list'" in result
```


Tests for q2.3 (points: 1)
```
>>> def test_type_error():
...     with pytest.raises(TypeError):
...         area('10')
>>> test_type_error()
```
