```python
# Initialize Otter
import otter
grader = otter.Notebook("worksheet5.ipynb")
```

# Worksheet 5 - NumPy

```python
import numpy as np
```

### Exercise 1
rubric={autograde:1}

Create the following array and save it as a variable named `ex1`:

```python
array([[1, 2],
       [3, 4],
       [5, 6]])
```

```python
ex1 = ...
```

```python
grader.check("ex1")
```

### Exercise 2

Create a 1D array of 5 zeros and save it as a variable named `ex2`:

```python
ex2 = ...
```

```python
grader.check("ex2")
```

### Exercise 3

Create a 1D array of 10 and save it as a variable named `ex3`:

```python
ex3 = ...
```

```python
grader.check("ex3")
```

### Exercise 4

Create a 1D array of five numbers, each having the value of 3.141, and save it as a variable named `ex4`:

```python
ex4 = ...
```

```python
grader.check("ex4")
```

### Exercise 5

Create a 1D array of the integers 1 to 20 and save it as a variable named `ex5`:

```python
ex5 = ...
```

```python
grader.check("ex5")
```

### Exercise 6

Create a 5 x 5 matrix of ones with a dtype `int` and save it as a variable named `ex6`:

```python
ex6 = ...
```

```python
grader.check("ex6")
```

### Exercise 7

Create an 3D matrix of 3 x 3 x 3 full of random numbers drawn from a standard normal distribution (hint: `np.random.randn()`) and save it as a variable named `ex7`:

> Hint: To ensure the code reproducibility, set the random seed = 42 using the `np.random.seed()` function.

```python
# set your seed to 42 for reproducibility
ex7 = ...
```

```python
grader.check("ex7")
```

### Exercise 8

Reshape the above array (`ex7`) into shape (27,) and save it as a variable named `ex8`:

```python
ex8 = ...
```

```python
grader.check("ex8")
```

### Exercise 9

Reshape the following `arr1` into shape `(6, 1)` and save it as a variable named `ex9`:

```python
arr1 = np.array([[1, 2], [3, 4], [5, 6]])
arr1
```

```python
ex9 = ...
```

```python
grader.check("ex9")
```

### Exercise 10

Create a `(6, 6)` array by multiplying the following `arr2` by itself (**hint**: transpose the array to leverage broadcasting here). The first element (top left corner) should be 1, and the last element (bottom right corner) should be 36. Save it as a variable named `ex10`.

```python
arr2 = np.arange(1, 7)[np.newaxis, :] # np.newaxis is an alias for None
arr2
```

```python
ex10 = ...
```

```python
grader.check("ex10")
```

### Exercise 11

"Center" the following `arr3` by subtracting the mean from every element, then flatten it to shape `(20,)` and use `numpy.sort` to sort it in descending order. Save it as a variable named `ex11`.

```python
arr3 = np.arange(1, 21).reshape(4, 5)
arr3
```

```python
ex11 = ...
```

```python
grader.check("ex11")
```

Congratulations! You are done the worksheet!!! Pat yourself on the back and submit your worksheet to Gradescope!

Tests for ex1 (points: 1)
```
>>> assert type(ex1) == np.ndarray, 'ex1 should be a numpy array'
```
```
>>> assert ex1.shape == (3, 2), 'ex1 should have shape (3, 2)'
```
```
>>> assert ex1[0][0] == 1
```
```
>>> assert ex1[0][1] == 2
```
```
>>> assert ex1[2][0] == 5
```
```
>>> assert ex1[2][1] == 6
```


Tests for ex10 (points: 1)
```
>>> assert type(ex10) == np.ndarray, 'ex10 should be a numpy array'
```
```
>>> assert ex10.shape == (6, 6), 'ex10 should have shape (6, 6)'
```
```
>>> assert ex10[0, 0] == 1
```
```
>>> assert ex10[-1, -1] == 36
```


Tests for ex11 (points: 1)
```
>>> assert type(ex11) == np.ndarray, 'ex11 should be a numpy array'
```
```
>>> assert ex11.shape == (20,), 'ex11 should have shape (20,)'
```
```
>>> assert ex11[0] == 9.5
```
```
>>> assert ex11[-1] == -9.5
```
```
>>> assert ex11.sum() == 0
```


Tests for ex2 (points: 1)
```
>>> assert type(ex2) == np.ndarray, 'ex2 should be a numpy array'
>>> assert ex2.shape == (5,), 'ex2 should be a 1D array of length 5'
>>> assert np.array_equal(ex2, np.array([0.0, 0.0, 0.0, 0.0, 0.0])), 'q2_1 should be a 1D array of zeros'
```


Tests for ex3 (points: 1)
```
>>> assert type(ex3) == np.ndarray, 'ex3 should be a numpy array'
>>> assert ex3.shape == (10,), 'ex3 should be a 1D array of length 10'
>>> assert np.array_equal(ex3, np.array([1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0])), 'q2_2 should be a 1D array of ones'
```


Tests for ex4 (points: 1)
```
>>> assert type(ex4) == np.ndarray, 'ex4 should be a numpy array'
>>> assert ex4.shape == (5,), 'ex4 should be a 1D array of length 5'
>>> assert np.array_equal(ex4, np.array([3.141, 3.141, 3.141, 3.141, 3.141])), 'ex4 should be a 1D array of 3.141s'
```


Tests for ex5 (points: 1)
```
>>> assert type(ex5) == np.ndarray, 'ex5 should be a numpy array'
>>> assert ex5.shape == (20,), 'ex5 should be a 1D array of length 20'
>>> assert np.array_equal(ex5, np.array([1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20]))
```


Tests for ex6 (points: 1)
```
>>> assert type(ex6) == np.ndarray, 'ex6 should be a numpy array'
>>> assert ex6.shape == (5, 5), 'ex6 should be a 2D array of shape (5, 5)'
>>> assert np.array_equal(ex6, np.array([[1, 1, 1, 1, 1], [1, 1, 1, 1, 1], [1, 1, 1, 1, 1], [1, 1, 1, 1, 1], [1, 1, 1, 1, 1]])), 'ex6 should be a 2D array of ones'
```


Tests for ex7 (points: 1)
```
>>> assert type(ex7) == np.ndarray, 'ex7 should be a numpy array'
>>> assert ex7.shape == (3, 3, 3), 'ex7 should be a 3D array of shape (3, 3, 3)'
```


Tests for ex8 (points: 1)
```
>>> assert type(ex8) == np.ndarray, 'ex8 should be a numpy array'
>>> assert ex8.shape == (27,), 'ex8 should be a 1D array of length 27'
```


Tests for ex9 (points: 1)
```
>>> assert type(ex9) == np.ndarray, 'ex9 should be a numpy array'
```
```
>>> assert ex9.shape == (6, 1), 'ex9 should have shape (6, 1)'
```
```
>>> assert ex9[0] == 1
```
```
>>> assert ex9[-1] == 6
```
